## Tensorflow Basic Usage Test -- mean

### Part A(2 Points)

Write code to compute the mean of a list of numbers

<!-- grade_id: mean_func, points: 0, grade: false, solution: true -->
```python
def mean(x):
    """ Write code to compute the mean of a list of numbers """
    ### BEGIN SOLUTION
    
    ### END SOLUTION
```

<!-- grade_id: test_mean, points: 2, grade: true, solution: false -->
```python
""" Check that the 'mean' function is correct """
assert(mean({1}) == 1)
assert(mean({1, 2}) == 1.5)
### BEGIN HIDDEN TESTS
assert(mean({5.5, 0, 2, 3.4}) == 2.725)
assert(mean(range(100)) == 49.5)
assert(mean(range(100, 0, -1)) == 50.5)
### END HIDDEN TESTS
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-4-cbb0d3970b91> in <module>()
      1 """ Check that the 'mean' function is correct """
----> 2 assert(mean({1}) == 1)
      3 assert(mean({1, 2}) == 1.5)
      4 assert(mean({5.5, 0, 2, 3.4}) == 2.725)
      5 assert(mean(range(100)) == 49.5)

AssertionError:
```

### Part B(3 Points)

Please describe the deference between an arithematic mean, a harmonic mean, and a geomatric mean.

<!-- grade_id: describe_mean, points: 3, grade: true, solution: true -->
Arithmatic mean:

Harmonic mean:

Geomatric mean:
